<subject>
Homework How-To
</subject>

<details>

**Naming Conventions**

When naming your files for upload, you must follow the format below:

    <uni>_<assignment>_<details [optional]>.<extension>

For example, if I needed to hand in HW 0, any of the below formats would be sufficient for a file name:

- pl2648_hw0.txt
- pl2648_hw0.sh
- pl2648_hw0_all_in_one.txt
- Pl2648_hw0_bash_program.sh

This naming format allows for autograding of all assignments. If your files are not named with this format, you should expect a grade of zero for the assignment.

**Grading**

Possible points on late homeworks are deducted by 50% for each day they are late. For example, if you get 80% of the total possible credits on a homework but hand in that homework a day late, you would get 40%. Assignments two days late get zero points.

Once solutions are posted and graded assignments are handed back, students have 1 week to bring their grading discrepancies to a CA for consideration of possible grading errors. 

**Getting Help**

Asking for help is a great way to increase your chance of success. However there are some rules. When asking for help (especially from a fellow student), *you can show your helper your code but you can not view theirs*. You work needs to be your own. 

If you need to reach out to a CA or the Professor for help, please do so via Piazza and not via email. Answers given via Piazza will help you as well as other students. Thus, emails will always have a lower priority for response than Piazza questions. If you do email the professor or TA, please make a note of what section you are in. This helps us identify you in courseworks faster.
</details>

**Q**: Build a function to return a regex expression. The expression, when passed to the appropriate `re` function, should find all of the dates of the format YYYY-MM-DD in the in following string.

```python
>>> string = """
2010-04-29 1984 ABC-123!
Feb 18th, 2018. 
888-27-8949 2048-07-07
"""

>>> pattern = get_pattern()
>>> func = get_function()
>>> func(pattern, string)
['2010-04-29', '2048-07-07']
```

[1 point]

<!-- grade_id: cell-191bb99330f39907, points: 0, grade: false, solution: true -->
```python
import re

def get_pattern():
    # YOUR CODE HERE
    raise NotImplementedError()

def get_function():
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-7c5fc12bf1bedd77, points: 0.25, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-37f7551c7035fb27, points: 0.5, grade: true, solution: false -->
```python
### BEGIN TESTS
string = """
2010-04-29 1984 ABC-123!
Feb 18th, 2018. 
888-27-8949 2048-07-07
12048-07-074
"""

pattern = get_pattern()
func = get_function()
result = func(pattern, string)

assert result == ['2010-04-29', '2048-07-07'], result
### END TESTS
```

<!-- grade_id: cell-7719cc19d23054ae, points: 0.25, grade: true, solution: false -->
```python

```

**Q**: Write a function that takes a string like the following and returns a match object if a match is possible. 

    '2018 Mar 24 02:41:00 AM'

The match object should have the following named groups:

    - year
    - month_name
    - day
    - hour
    - minute
    - second
    - meridiem
    
For example, the output of ...

```python
match = pattern.match('2018 Mar 24 02:41:00 AM')
print('MATCH', match, sep=': ')
print('GROUPS', match.groups(), sep=': ')

for name in ('year', 'month_name', 'day', 'hour', 'minute', 'second', 'meridiem'):
    print(name.upper(), match.group(name), sep=': ')
 ```
where `pattern` is the compiled regex expression should be:

    MATCH: <_sre.SRE_Match object; span=(0, 23), match='2018 Mar 24 02:41:00 AM'>
    GROUPS: ('2018', 'Mar', '24', '02', '41', '00', 'AM')
    YEAR: 2018
    MONTH_NAME: Mar
    DAY: 24
    HOUR: 02
    MINUTE: 41
    SECOND: 00
    MERIDIEM: AM
    

[2 points]

<!-- grade_id: cell-be554692bfc729a9, points: 0, grade: false, solution: true -->
```python
import re

def get_match_object(string):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-20885ead6e5a1fac, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('year') == '2018'
### END TESTS
```

<!-- grade_id: cell-d65adb32954b254a, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('month_name') == 'Mar'
### END TESTS
```

<!-- grade_id: cell-64860e1dd88fbc1d, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('day') == '24'
### END TESTS
```

<!-- grade_id: cell-47661529f963e201, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('hour') == '02'
### END TESTS
```

<!-- grade_id: cell-2e203308d6349075, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('minute') == '41'
### END TESTS
```

<!-- grade_id: cell-51337e5d0dacb4ad, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('second') == '00'
### END TESTS
```

<!-- grade_id: cell-d49f5d2076486460, points: 0.25, grade: true, solution: false -->
```python
### BEGIN TESTS
match = get_match_object('2018 Mar 24 02:41:00 AM')
assert match.group('meridiem') == 'AM'
### END TESTS
```

<!-- grade_id: cell-d5f2c2c1e38f1be6, points: 0.25, grade: true, solution: false -->
```python

```

**Q**: NYC publishes 311 call reports for public review. A specific slice of that data has been made available at the following link:

https://data.cityofnewyork.us/Social-Services/311-11-19-2018-/6ir3-7kwf/data

Download the CSV of this data.

Using the `re` module, find the number at the beginning of each "Incident Address". 

For "17 WEST 74 STREET" the number would be 17, for "88-28 74 PLACE" the number would be 88-28. Ignore addresses that are empty strings. If the street number is 123A, consider the number to be 123. 

If the number is a hyphenated number, peform the subtraction (ie. 88-28 = 60). Sum up all of the results of the hyphenated addresses. (ie. 60 + (24-11) + ...)

Sum up all the un-hyphenated addresses.

Return the pair as a tuple from the function provided. The function should take the absolute path to the CSV file.

[2 points]

<!-- grade_id: cell-0b8c920f2ed2ebfc, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()


def get_address_sums(path: str) -> (int, int):
    # YOUR CODE HERE
    raise NotImplementedError()
    return (unhyphenated, hyphenated)
```

<!-- grade_id: cell-695afe2438cf8f12, points: 0, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-7ab2046a0323a2a7, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-3d7a0349af2c6f3f, points: 1, grade: true, solution: false -->
```python

```